```python
# Initialize Otter
import otter
grader = otter.Notebook("HW4.ipynb")
```

## <u>Name:</u> [Your name] and (your onyen)

Names and Onyens of fellow students you discussed Homework4 problems and ideas:

- Students: 

We encourage discussing ideas and brainstorming with your peers, but the final text, code, and comments in this homework assignment MUST be 100% written by you as mentioned in syllabus.



# Homework 4: Probability, Sampling, Correlation and Causation  (50 points)

## Problem 1. Probability (10 points)

In this problem, you'll be analyzing the fairness of a 6-sided die.

A street vendor asks you if you'd like to play a game.  The game is:
* pay \$10 to roll a 6-sided die 3 times
* if any of those rolls is a 6 he pays you $11.  

**Problem 1.1 (1 points).** What are your chances of rolling at least one 6 in a game? (show how you calculate this)

*Write your answers here*

**Problem 1.2 (1 points).** What are your chances of rolling only one 6 in a game? (show how you calculate this)

*Write your answers here*

**Problem 1.3 (1 points).** What are your chances of winning the game given that the first two rolls are not a 6? (show how you calculate this)

*Write your answers here*

**Problem 1.4 (1 points).** What are your chances of winning the game given that the first roll is not a 6? (show how you calculate this)

*Write your answers here*

**Problem 1.5 (1 points).** What are your chances of losing the game given that the first two rolls are not a 6? (show how you calculate this)

*Write your answers here*

**Problem 1.6 (1 points).** What are your chances of losing the game given that the first roll is not a 6? (show how you calculate this)

*Write your answers here*

**Problem 1.7 (1 points).** If you play this game four times, what is the probability that you lose all four games?

*Write your answers here*

**Problem 1.8 (1 points).** What is the **expected value** of your profit or loss if you play the game once?

*Write your answers here*

**Problem 1.9 (1 points).** What is the **expected value** of your profit or loss if you play the game 10 times?

*Write your answers here*

**Problem 1.10 (1 points).** If you play the game 100 times, are you guaranteed to win at least one of the games? If not, what is the probability?

*Write your answers here*

## Problem 2. Just how old are you? (15 points) 
Suppose you are asked to provide information about the **age of permanent residents of North Carolina**, so you decide to obtain age data for **100 people**.

**Problem 2.1 (3 points).**
Initially, you ask the age of the **first 100 people you meet on campus**.  
1. What is this type of sampling called? Pick **one** closest and most specific answer.
3. What are potential biases that could result from this sampling approach?  Name **at least two**.  
4. Would you expect this to be a representative sample of North Carolina residents?

*Write your answers here*

**Problem 2.2 (3 points).** 
You decide to look up the **proportions of race/ethnicity** for North Carolina permanent residents. You calculate the corresponding **proportions of 100**, then collect data again by wandering campus and asking the age of exactly that many people of each race/ethnicity.  

For example, the [2020 US Census](https://www.census.gov/library/stories/state-by-state/north-carolina-population-change-between-census-decade.html) says the four largest NC race and ethnicity groups are White alone 62%, Black alone 12%, Hispanic 19%, and Asian alone 6%. So you would ask the age of the first 62 White people you meet, the first 12 Black people you meet, and so on.

1. What type of sampling is this? Pick **one** closest and most specific answer.
2. Would you expect it to be more or less representative than your previous data from Problem 1.1?  
3. What are the most significant biases that would remain?  Name **at least two**.

*Write your answers here*

**Problem 2.3 (3 points).** 
Alternatively, suppose you get access to full census data for all residents of the state of North Carolina that includes age as an attribute.  You take the age of the **first 100 rows of the table** and drop the rest.  What potential concerns do you have with this approach?

*Write your answers here*

**Problem 2.4 (2 points).** 
Finally, you decide to take **100 random samples from the census data** table.
1. What type of sampling is this? Pick **one** closest and most specific answer.
2. Would you expect this to be a representative sample of North Carolina residents?
3. Do you think it matters if you allow replacement for this sampling?  Why or why not?
4. If the census data is in a table named `nc_people` that has a row for every person, write code that uses pandas to randomly sample 100 people from `nc_people` and put them in a new table named `nc_people_100`.

*Write your answers here*

**Problem 2.5 (1 point).** 
Which of the following is an example of sampling?
1. Enumerating all the possible outcomes of drawing 1 card from a deck
2. Counting the number of planets in our solar system
3. Surveying a piece of property to mark its boundaries
4. Asking people in your class who sit next to you their favorite Pokemon 

**Problem 2.6 (3 points)** Let's now try out this idea on Penguins data. 

```python
# Run this cell to load penguins table
import pandas as pd
penguins = pd.read_csv('~/110-F25/demos/palmer_penguins.csv')
penguins.shape
```

Say we sample 10% of the penguins to estimate characteristics of the population (344 penguins).

(1 point) Write code that calculates the **average Body Mass** and **Sex distribution** of the three groups:

    1. First 10% of the penguins listed in the table.
    2. Randomly sampled 10% of the penguins.
    3. All the penguins.

```python
# Your code here. Add new cells as needed

...
```

```python
# Your code here. Add new cells as needed
...
```

(2 points) For **average Body Mass**, which sample gave a closer estimate to the population (344 penguins) mean? Why do you think that happened? Feel free to inspect the table further as needed.

How about for **Sex distribution**? Was this answer different or same from earlier? Why do you think that happened?

[Write your answer here]

## Problem 3. Modeling (13 points)
This problem deals with determining whether a proposed model is reasonable for describing one attribute as a function of another for ROUSes.

First, you will be evaluating which of 3 models would best describe how `Length` is a function of `Age`.

Start by running the code below.

```python
import pandas as pd
import seaborn as sns
rouses_model_df = pd.read_csv("ROUSes.csv")
```

**Problem 3.1 (3 points).** Add a new feature (column) to `rouses_model_df` for each of the 3 models below, and fill in the feature with data generated by the model.  Name the features `'Length Model 1'`, `'Length Model 2'`, and `'Length Model 3'`.  Note that the first model has a square root; for this you can use `**0.5`.  

**Hint**: if you are unsure of mathematical operation ordering, parentheses are highly recommended!

  Length Model 1 = 1.08 * $\sqrt{Age}$

  Length Model 2 = 1.0 + 0.128 * $Age$

  Length Model 3 = 1.0 + 0.00229 * $Age^2$



```python
# write your code here, replacing ...
...
```

```python
grader.check("Q3.1")
```

**Problem 3.2 (3 points).** There are many ways to compare data generated from a model to the actual data.  One of the easiest is to overlay plots of the actual and generated data.  One very simple way to do this is to put multiple scatterplots in the same cell, one after the other; seaborn is smart enough to automatically change the color for each of the plots.  Do this in the below; your plot should look like this: \
 ![](LengthModel.png)

**Hint**: https://seaborn.pydata.org/tutorial/properties.html#marker


```python
# write your code here, replacing ...
...
```

**Problem 3.3 (2 points).**  Just by looking at the plots, which of the models would you say is the best fit to the actual data? Why (1~3 sentences)? 

[Type your answer here, replacing this text.]

**Problem 3.4 (5 points).** Using any technique you wish, decide which of the models below is the best model for `Weight`.  Show some kind of work or visualization in the Python cell, and write what model matches best (and Why!!) in the cell below that.

  Weight Model 1 = 51 * $\sqrt{Age}$

  Weight Model 2 = 6.9 * $Age$

  Weight Model 3 = 0.125 * $Age^2$

Autograder portion is 2/5 points.

```python
# show your work here
...
```

```python
grader.check("Q3.4")
```

_Type your answer here, replacing this text._

## Problem 4. Causality vs. Correlation (12 points)
An important part of data science is understanding the relationships between data.  **Correlation** is the degree to which data change with one another.  **Causality** implies a direct relationship where changes in one thing directly results in changes to another thing.  Data can be correlated without any causality relationship; in fact, this occurs often when the data both depend on a third variable.  For example, number of shark attacks and ice cream sales are often correlated, but one does not cause the other; they both depend on weather conditions.  

**Problem 4.1 (2 points).** Run the code below and observe the plot.  Does the ***y*** value appear to be correlated with the ***x*** value? Why or why not? Calculate its correlation coefficient.


```python
import seaborn as sns
x_data = [0,1,2,3,4,5,6,7,8,9,10]
y_data = [1,4,7,10,2,5,8,9,6,3,0]
sns.scatterplot(x=x_data, y=y_data)
```

```python
# Your code here
```

_Type your answer here, replacing this text._

**Problem 4.2 (2 points).** Run the code below and observe the plot.  Does the ***y*** value appear to be correlated with the ***x*** value? Why or why not? Calculate its correlation coefficient.


```python
x_data = [0,1,2,3,4,5,6,7,8,9,10]
y_data = [10,13,14,16,18,20,21,24,26,27,30]
sns.scatterplot(x=x_data, y=y_data)
```

```python
# Your code here
```

_Type your answer here, replacing this text._

**Problem 4.3 (8 points).** For each of the following cases, decide whether it is an example of **causality**, **correlation**, or **neither**.  For each answer, explain why in a single sentence.

* Case A. Roulette wheel spins and crap table dice rolls.

_Type your answer here, replacing this text._

* Case B.
  
 ![](this_is_fine.jpg)

_Type your answer here, replacing this text._

* Case C. Golf club swing speed and golf ball initial flight speed.

_Type your answer here, replacing this text._

* Case D. Shoe size and reading ability.

_Type your answer here, replacing this text._

## Submission to Gradescope

### Step 1 – Run All Cells and Save
- My favorite thing to do in Jupyter Notebook is to click the ⏩️ button on the toolbar. This refreshes the memory, then runs all cells in order.
- Then, click **File → Save Notebook** to make sure your notebook file (`.ipynb`) is up to date. Or run `command (or control) + s`.

### Step 2 – Export to PDF
  1. Choose **File → Print**.
  2. In your browser’s print dialog, change the **Destination** to “Save as PDF.”
  3. Click **Save** and name the file clearly (e.g., `lab1.pdf`).

### Step 3 - Run the Last Cell
1. Run the last cell (below)
2. Download the zip file to your computer.

### Step 4 – Submit to Gradescope
- Log into **Gradescope** and open the assignment. You can get there from Canvas assignment as well.
- Upload **both files** by highlighting both items and draggin and dropping them to the submission box:
  1. Your **`.zip` file** (generated by otter).
  2. Your **`.pdf` file** (exported from the notebook).
- Make sure your submission is **complete and readable** before clicking submit.

### Step 5 – Confirm Your Submission
- After submitting, check Gradescope’s preview to ensure your files display properly.
- If anything is missing or unreadable, re-export and resubmit before the deadline.


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope HW2

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for Q3.1 (points: 3)
```
>>> assert rouses_model_df['Length Model 1'].equals(1.08 * rouses_model_df['Age'] ** 0.5)
>>> assert rouses_model_df['Length Model 2'].equals(1 + 0.128 * rouses_model_df['Age'])
>>> assert rouses_model_df['Length Model 3'].equals(1.0 + 0.00229 * rouses_model_df['Age'] ** 2)
```


Tests for Q3.4 (points: 2)
```
>>> assert rouses_model_df['Weight Model 1'].equals(51 * rouses_model_df['Age'] ** 0.5)
>>> assert rouses_model_df['Weight Model 2'].equals(6.9 * rouses_model_df['Age'])
>>> assert rouses_model_df['Weight Model 3'].equals(0.125 * rouses_model_df['Age'] ** 2)
```
