```python
# Initialize Otter
import otter
grader = otter.Notebook("lab2.ipynb")
```

# Welcome to Lab 2

- Your name: [Your name here]
- Your onyen: [Your onyen here] 

Collaborators

-
-
-


# 1. Data collection (10 minutes)

**Q.** We'll start today by collecting survey data from 4 classmates + yourself. We'll represent our collected data as a table, which you can write in Markdown as below. Here's a quick guide to Markdown tables (https://www.codecademy.com/resources/docs/markdown/tables), and you can also double click on the cell below to see how I wrote it.

| # | First name | What vegetable would you be? | Shoe size | Birthday month | (Question of your choice) |
|---| -----------| -----------| -----------| -----------| -----------| 
|1 | | | | | 
|2 | | | | | 
|3 | | | | | 
|4 | | | | | 
|5 | | | | | 


Please mingle (try to talk to someone you haven't talked to yet!) and fill out the table. For example, here's one row of data that I have gathered at home:

| # | First name | What vegetable would you be? | Shoe size | Birthday month | Are you a cat? |
|---| -----------| -----------| -----------| -----------| -----------| 
|1 | Puppycat | Corn | 0 | July | True |

If you are working on this outside of lab time (e.g. excused absence, late add from waitlist), just ask people in your dorm or text some friends.

**Q.** Fill in the blank with either `row` or `column`:

- Each *classmate* corresponds to one _____.
- Each *question* corresponds to one _____.

# 2. Data types (5 minutes)

In last lecture, we learned four different data types in Python: `int`, `float`, `str`, and `bool`. If you wanted to work with these data using Python, which types would be the best fit and why? Would you use the same type for all of them? Discuss with your group.

**Q.** Fill in your guess for the best data type for each piece of information. 

- First name:
- Vegetable:
- Shoe size:
- Birthday month:
- Cat (replace with your question):


Discuss your answers with your group. There can be multiple valid answers. For example, `'July'` or `7` are both valid ways to talks about birthday months. 

# 3. TA-led discussion on data types and units (10 minutes)

**Q.** Add a new Markdown cell below and write quick 2-3 sentences about what the class discussed (Imagine you're summarizing for a friend who missed recitation today).

# 4. Name (variable) assignment (10 minutes)

**Q.** Now that you have made a decision on data types, let's practice using it. You can go back to your small group to work on this together.

1. Pick one row from your table.
2. In the cell below, assign the first name of the individual to the variable `name`. For example, `name = 'Puppycat'`.
3. In the next cell below, check the type of the variable `name`. Do you remember what function to call?
4. Repeat steps 2-3 with the remaining four columns. Add new cells as needed. Try to pay attention to the data type as well as what a descriptive variable name would be.

```python
# Write your code here
name = ...
```

```python
# Write your code for other variables here

```

```python
grader.check("name")
```

# 5. Exercises (15 minutes)
Here are some exercises that are similar to questions in HW1. Especially when you're brand new to programming, this is great practice when you have groupmates and TA to help.

**Q.** Variables and printing
- Assign `pet = "dog"` and print it.
- Then reassign `pet = "cat" + "dog"` and print again.

```python
# Your code here
```

Make sure you understand what those lines did. No need to provide written answer here.

**Q.** Predict boolean logic results

Without running, predict `True`/`False` for each line. Then run to check yourself.
Replace "predict then check" with the answer for each line.

```python
x = 7
print(x > 5)      # predict then check
print(x == 10)   # predict then check
print(x != 3)    # predict then check
print(x < 3 or x > 6)    # predict then check
print(x > 3 and x < 6)    # predict then check
print(not x > 6)    # predict then check
```

**Q.** Simple `if`

Run the code. Then edit the value of `temp` to 60, -10, 40 and observe which message prints.

```python
temp = 32  # try changing this value
if temp < 40:
    print("It's cold!")
else:
    print("Not so cold")
```

What is the printed message? 
- temp is 32: [Your answer here]
- 60: [Your answer here] 
- -10:[Your answer here]
- 40: [Your answer here]

**Q.** Built-in functions and string methods

Run the next few cells and talk with your group to understand what the functions and string methods do to the variable `word='banana'`.

```python
word = 'banana'

len(word)
```

```python
word = 'banana'

print('1. word is:', word)

w = word.upper()
print('2. w is:', w)
print('3. word is:', word)
```

```python
word = 'banana'

print('1. word is:', word)

word = word.upper()
print('2. w is:', w)
print('3. word is:', word)
```

What is the difference between the two cells above? Why do different behaviors occur?

[Your answer here]

```python
word.lower()
```

```python
word = 'bAnana'

word.replace("a", "o")
```

**Q.** There are more string methods listed in https://www.w3schools.com/python/python_ref_string.asp. Now, if you have a filename such as `"report.csv"`, which of the string methods would you use to check whether the filename ends with `'.csv'`? Click on that method from the table, take a look at their example, then try to write code that:

1. Assigns variable `ends_ok` to whether the filename ends with `'.csv'` or not.
2. If `ends_ok`, then print "valid".
3. If not, then print "invalid".


```python
filename = "report.csv" # try different values

# Your code here
```

**Q.** Add `elif`

Classify a number into three buckets. Try different values for n and talk through the logic with your group. No text answer needed.

```python
n = 15  # try different values

if n >= 20:
    print("large")
elif n >= 5:
    print("medium")
else:
    print("small")

if n == 13: # How is this different from elif? Try changing n to 13 in the first line.
    print("thirteen")
```

**Q.** Translating question to more complex logic and control

Print `Special` if `temp < 32` or `wind > 20`; otherwise print `Normal`. Then print `Not California` if `state` is not California, else `California spotted!`.

```python
temp = 28
wind = 10
state = "Oregon"

# Your code here
```

**Q.** You can assigning boolean values to variables, then using it in if/else statements.

```python
x = 12 # Try different values
y = 20 

# Assign boolean values
is_x_less = (x < y) # What is the type of this variable? [Your answer here]
is_y_big = (y > 15) 

print(is_x_less)  # What do you expect to see? [Your answer here]
print(is_y_big)  # What do you expect to see? [Your answer here]
```

```python
if is_x_less and is_y_big:
    print("x is smaller than y, and y is big!")
else:
    print("Condition not met.")
```

How would you combine `is_x_less` and `is_y_big` to create this condition?

    x is at least as big as y, and y is <= 15.

```python
# your code here
```

## Submission to Gradescope

### Step 1 – Run All Cells and Save
- My favorite thing to do in Jupyter Notebook is to click the ⏩️ button on the toolbar. This refreshes the memory, then runs all cells in order.
- Then, click **File → Save Notebook** to make sure your notebook file (`.ipynb`) is up to date. Or run `command (or control) + s`.

### Step 2 – Export to PDF
  1. Choose **File → Print**.
  2. In your browser’s print dialog, change the **Destination** to “Save as PDF.”
  3. Click **Save** and name the file clearly (e.g., `lab1.pdf`).

### Step 3 - Run the Last Cell
1. Run the last cell (below)
2. Download the zip file to your computer.

### Step 4 – Submit to Gradescope
- Log into **Gradescope** and open the assignment. You can get there from Canvas assignment as well.
- Upload **both files** by highlighting both items and draggin and dropping them to the submission box:
  1. Your **`.zip` file** (generated by otter).
  2. Your **`.pdf` file** (exported from the notebook).
- Make sure your submission is **complete and readable** before clicking submit.

### Step 5 – Confirm Your Submission
- After submitting, check Gradescope’s preview to ensure your files display properly.
- If anything is missing or unreadable, re-export and resubmit before the deadline.


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope Lab2

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for name (points: 0)
```
>>> assert type(name) == str
```
