```python
# Initialize Otter
import otter
grader = otter.Notebook("HW5.ipynb")
```

# Homework 5: Inference and Hypothesis Testing (50 points)

## <u>Name:</u> [Your name] and (your onyen)

Names and Onyens of fellow students you discussed Homework 5 problems and ideas:

- Students: 

We encourage discussing ideas and brainstorming with your peers, but the final text, code, and comments in this homework assignment MUST be 100% written by you as mentioned in syllabus.



## Problem 1. Getting started on the project proposal (4 points)

This question is here to make sure you don't all work on your project proposal a day before. :)))) 

- You must pick a dataset as a group. There's no penalty for changing it down the road, but you might as well pick one now!

[Write your answer here]

- Conduct individual data exploration: calculate **summary statistics** and create **at least two** basic visualizations of your dataset.
  
Try to create visualizations that are different from those created by your teammates. 

```python
# To load dataset from a different folder, change the variable fn appropriately.
import pandas as pd
fn = 'palmer_penguins'
data = pd.read_csv(f'~/110-F25/project/{fn}.csv')
```

```python
# Write your code here
```

```python
# Write your code here
```

```python
# Write your code here
```

- What did you notice (3~4 sentences)?

[Write your answer here]

## Problem 2. When there are many samples (22 points)

In HW4, we sampled 10% of penguins to demonstrate good sampling practices. In this problem, we'll use the same dataset to demonstrate Law of Large Numbers and Central Limit Theorem. 

```python
# Run this cell
import pandas as pd
penguins = pd.read_csv('~/110-F25/demos/palmer_penguins.csv')
penguins.shape
```

### Problem 2.1 (2 points) Population and sample.
Here is our fictional situation. There are exactly 344 penguins left in the world (RIP). You're a researcher camping out in Antarctica, hoping to meet penguins and take their measurements. Due to logistical difficulties, you won't be able to gather all 344 of them in one place and record their characteristics. Select the most appropriate choices to complete the following sentences about this fictional situation. Make edits directly to the text.

- 344 penguins are (the population of interest) or (the sample)
- The average weight of the 344 penguins is (population mean) or (sample mean)
- The average weight of 34 out of 344 penguins is (a parameter) or (a statistic)

### Problem 2.2 (2 points) Population mean and standard deviation.

Caculate the **population mean** and **population standard deviation** of **Body Mass**. Assign them to variables `pop_mean` and `pop_std`.

```python
...
pop_mean, pop_std
```

```python
grader.check("Q2.2")
```

### Problem 2.3 (6 points) Law of Large Numbers

In lecture, we simulated a die roll by sampling from uniform distribution. Here, we will be sampling rows from the penguins table.

**(1 points)** Should these rows be sampled with replacement (`replace=True`) or without replacement (`replace=False`) to get closer to the assumptions of LLN? Why?

[Write your answer here]

**(3 points)** Complete the code below that demostrates LLN on penguins data. Similar to lecture demo, instead of sampling one row at a time inside the for loop, let's:
1. sample `max_k` of them all at once, then
2. inside the for loop as `k` changes,
3. take the mean of first `k` samples.

Make sure to have a horizontal line indicating the population mean.

Take a look at lecture slides and demo code for help.

```python
import matplotlib.pyplot as plt
import seaborn as sns

max_k = 4000  
# sample max_k penguins
max_k_penguins = ...
all_means = []

for k in range(1, max_k):
    # first k out of max_k penguins
    sampled_penguins = ...
    # calculate sample mean
    sample_mean = ...
    all_means.append(sample_mean) # save sample mean

fig, ax = plt.subplots(1)
sns.scatterplot(x=range(1, max_k), y=all_means, label='sample mean', s=2)
ax.set_xlabel('Number of samples')
ax.set_ylabel('Mean of observed data')
# horizontal line code goes here, replacing ...
...
ax.legend()
```

**(2 points).** Why is it enough to measure *some* penguins instead of all 344? How many penguins do you think you'd need to see before your estimate is "good enough"? Why?

[Write your answer here]

### Problem 2.4 (12 points) Central Limit Theorem

**(3 points)** What we explored in the LLN problem was the idea of sampling Penguins during *one* Antarctica trip. What happens if we made 2, 4, 10, 100s of Antarctica experiments? What can we say about the sample means from all of the different trips?

Complete the code below that demostrates CLT on penguins data. It is very similar to lecture demo. Your `num_exp` and `k` can be set to any reasonable number at time of submission.

```python
import matplotlib.pyplot as plt
import seaborn as sns

num_exp = 800 # number of times you repeat the experiment

k = 80 # number of cute penguins you meet

all_means = []

for exp in range(num_exp): 
    # sample k penguins
    sampled_penguins = ...
    # calculate sample mean
    sample_mean = ...
    all_means.append(sample_mean)

fig, ax = plt.subplots(1)
sns.histplot(all_means, label='sample means', stat='density')
ax.set_xlabel('Sample mean of body mass')
# horizontal line code goes here, replacing ...
...
ax.legend()

#### theoretical distribution according to CLT ####
from scipy.stats import norm
import numpy as np

sigma_sampling = pop_std / np.sqrt(k)

xs = np.linspace(min(all_means), max(all_means), 200)
ys = norm.pdf(xs, pop_mean, sigma_sampling)
ax.plot(xs, ys, 'r-', lw=2, label='CLT approx')
ax.legend()
########################################

print('num_experiment:', num_experiment)
print('k (number of penguins):', k)
print('mean of sample means:', np.array(all_means).mean().round(1))
print('standard deviation of sample means:', np.array(all_means).std().round(1))
```

**(2 points) Varying k** 

With `num_exp=200` fixed, run the code with different values of `k`, then fill out this table. The numbers are calculated for you in the last two print statements.

|k   |  mean (of sample means) | std (of sample means) |
| -- | ----------------------- | --------------------- |
| 20 |  |  |
| 80 |  |  |
| 320 |  |   |

What trends did you notice? Try to be as detailed as possible.

[Write your answer here]

**(2 points) Modeling**

I propose three models that might fit our observation of std (of sample means): 

$$1. std \propto k^{-1}$$
$$2. std \propto k^{-1/2}$$
$$3. std \propto k^{-2}$$

Using your previously calculated `pop_mean` value, fill in the table with what `std` values would look like according to the different models. This should feel similar to the ROUSE Age and Weight modeling exercises we did in the previous homework.

|k   |  1. pop_mean/k | 2. pop_mean/sqrt(k) | 3. pop_mean/k^2|
| -- | ------------| ---------------- | ----------- |
| 20 |  |  |   | 
| 80 |  |  |  | 
| 320 |  |   |  | 

Which model fits the data best and why?

[Write your answer here]

**(2 points) Varying num_exp** 

This time with `k=80` fixed, run the code with different values of `num_exp`, then fill out this table. 

|num_exp   |  mean (of sample means) | std (of sample means) |
| -------- | ----------------------- | --------------------- |
| 200 |  | |
| 400 |  |  |
| 800 |  |  |

What trends did you notice?

[Write your answer here]

**(1 points)** Make a **histogram** of the body mass of all penguins with appropriate labels. This should be the 344 population penguins, not the sampled penguins.

```python
# Your code here
```

**(2 points) Shape.**  

Does this population distribution look **skewed** or **symmetrical**? Does this population distribution look like **Gaussian**? 

How does it compare with the shape of the **sample mean distribution** we plotted before? Why is that okay or not okay?

[Write your answer here]

## Problem 3. Loaded question (24 points)

In this problem, you'll be analyzing the fairness of a 6-sided die.

A street vendor asks you if you'd like to play a game.  The game is:
* pay \$10 to roll a 6-sided die 3 times
* if any of those rolls is a 6 he pays you $11.  

### Problem 3.1 (4 points) Hypothesis Testing.
You decide to play the game 20 times, and out of your 60 total rolls of the die, you roll a 6 **four times**.  You start to suspect something is not right, and you wonder if perhaps the die is NOT fair.  In particular, you feel that there might be **less than fair chance** of rolling a 6, even though the street vendor insists it is fair. You decide to conduct a hypothesis test.

1) State the null hypothesis for this situation.  Remember, you'll need to be able to "simulate under the null".
2) State the alternative hypothesis.
3) Which of the following statistics would you be appropriate for evaluating these hypotheses?
* a) cost to play the game
* b) number of 5's rolled
* c) proportion of 6's rolled
* d) total number of rolls
4) What is the value of this statistic for the 20 games you played?  Assign it to `measured_statistic1` in the provided cell. (1 point autograder)

*Write your answers to 1-4 here*

```python
# Assign the value of the statistic for the 20 games here
measured_statistic1 = ...
```

```python
grader.check("Q3.1.4")
```

### Problem 3.2 (10 points) Simulating under the null.
In this problem, you'll be writing code to simulate the test statistic for 60 rolls, and you'll run that simulation 1000 times. 

#### Problem 3.2.1 (2 points).
First, write a code that returns a `list` named `rolls` containing simulated 6-sided die rolls for 60 rolls. Import whatever Python package you need to do this.

```python
# write your code here
...
rolls = ...
# This is sampling from uniform distribution
rolls[:10]
```

```python
grader.check("Q3.2.1")
```

#### Problem 3.2.2 (2 points).
Next, write code to simulate 1000 times the test statistic for 60 die rolls, and store the test statistic in a table named `statistic1`. 

```python
import pandas as pd
statistic1 = pd.DataFrame(columns=['test statistic'])
for simulation in range(1000):
    # finish the code below
    simulated_rolls = pd.DataFrame(random.choices([1,2,3,4,5,6], k=60), columns=['roll']) # create a 1-column table of randomly simulated rolls
    # sixes is an array of True or False. True if roll is 6, False otherwise.
    sixes = ...
    # calculate the statistic for this simulated data using sixes
    sim_t1 = ...
    statistic1.loc[len(statistic1)] = [sim_t1] # add simulated result to end of table
statistic1
```

```python
grader.check("Q3.2.2")
```

#### Problem 3.2.3 (2 points).
Plot the empirical distribution of your statistics simulated under the null.  Plot a vertical line for the value of the measured statistic for your actual games.  Label your plot using the parameter `label=`.

```python
# 3. Plot empirical distribution of t1 (under the null)
import matplotlib.pyplot as plt
fig, ax = plt.subplots(1)
# write your code here
...
ax.legend()
```

#### Problem 3.2.4 (1 point).
Looking at the shape of your distribution, what type of distribution is your simulation data, approximately? We're referring to the distribution of the statistics simulated under the null.

*Write your answer here*

#### Problem 3.2.5 (2 points).
Write code to compute the p-value from the `statistic1` table.

```python
p_value1 = ...
p_value1
```

```python
grader.check("Q3.2.5")
```

#### Problem 3.2.6 (1 point).
Does your simulation provide evidence against the null?  (yes or no)  If so, is the evidence statistically significant?  Is it extremely statistically significant?  Explain the convention you are using to answer these questions.

*Write your answers here*

### Problem 3.3 (4 points) Hypothesis testing revisited.
Instead of investigating if there is a less than fair chance of rolling a 6 like you just did, now evaluate if there is an **unfair chance** of rolling a 6 (either more likely or less likely than fair).

1) State the null hypothesis for this situation.
2) State the alternative hypothesis.
3) Which of the following statistics would you be appropriate for evaluating these hypotheses?
* a) number of 6's rolled in a row
* b) abs(proportion of 6's rolled - 1/6)
* c) abs(proportion of 6's rolled + 1/6)
* d) abs(total number of rolls)
4) What is the value of this statistic for the 20 games you played?  Assign it to `measured_statistic2` in the provided cell.

*Write your answers to 1-4 here*

```python
# Assign your measured statistic here
measured_statistic2 = ...
```

```python
grader.check("Q3.3")
```

### Problem 3.4 (6 points) Simulating under the null revisited

#### Problem 3.4.1 (2 points).
Now, write code to simulate 1000 times this new test statistic for 60 die rolls, and store the results in a table named `statistic2`.  You can use the same `sample_die()` function you wrote before.

```python
statistic2 = pd.DataFrame(columns=['test statistic'])

for simulation in range(1000):
    # finish the code below
    simulated_rolls = ...
    sixes = ...
    sim_t2 = ...
    statistic2.loc[len(statistic2)] = [sim_t2] # add simulated result to end of table
statistic2
```

```python
grader.check("Q3.4.1.")
```

#### Problem 3.4.2 (2 points).
Plot the empirical distribution of your statistics simulated under the null.  Plot a vertical line for the value of the measured statistic for your actual games.  Label your plot using the parameter `label=`.

```python
fig, ax = plt.subplots(1)
# write your code here
...
ax.legend()
```

#### Problem 3.4.3 (1 points).
Write code to compute the p-value from the `statistic2` table.


```python
p_value2 = ...
p_value2
```

```python
grader.check("Q3.4.3")
```

#### Problem 3.4.4 (1 points).
Does your simulation provide evidence against the null?  (yes or no)  If so, is the evidence statistically significant?  Is it extremely statistically significant?  

*Write your answers here*

## Submission to Gradescope

### Step 1 – Run All Cells and Save
- My favorite thing to do in Jupyter Notebook is to click the ⏩️ button on the toolbar. This refreshes the memory, then runs all cells in order.
- Then, click **File → Save Notebook** to make sure your notebook file (`.ipynb`) is up to date. Or run `command (or control) + s`.

### Step 2 – Export to PDF
  1. Choose **File → Print**.
  2. In your browser’s print dialog, change the **Destination** to “Save as PDF.”
  3. Click **Save** and name the file clearly (e.g., `lab1.pdf`).

### Step 3 - Run the Last Cell
1. Run the last cell (below)
2. Download the zip file to your computer.

### Step 4 – Submit to Gradescope
- Log into **Gradescope** and open the assignment. You can get there from Canvas assignment as well.
- Upload **both files** by highlighting both items and draggin and dropping them to the submission box:
  1. Your **`.zip` file** (generated by otter).
  2. Your **`.pdf` file** (exported from the notebook).
- Make sure your submission is **complete and readable** before clicking submit.

### Step 5 – Confirm Your Submission
- After submitting, check Gradescope’s preview to ensure your files display properly.
- If anything is missing or unreadable, re-export and resubmit before the deadline.


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit PDF to Gradescope Homework5

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for Q2.2 (points: 2)
```
>>> assert pop_mean == penguins['Body Mass (g)'].mean()
>>> assert pop_std == penguins['Body Mass (g)'].std()
```


Tests for Q3.1.4 (points: 1)
```
>>> assert measured_statistic1 == 4 / 60
```


Tests for Q3.2.1 (points: 2)
```
>>> assert type(rolls) == list
>>> assert all([type(r) == int for r in rolls])
>>> assert len(rolls) == 60
>>> assert all([r <= 6 for r in rolls])
>>> assert all([r >= 1 for r in rolls])
```


Tests for Q3.2.2 (points: 2)
```
>>> assert len(statistic1) == 1000
>>> assert sum(statistic1['test statistic'] >= 0) == 1000
>>> assert sum(statistic1['test statistic'] <= 1) == 1000
>>> assert sixes.equals(simulated_rolls['roll'] == 6)
>>> assert sim_t1 == sum(sixes) / len(simulated_rolls)
```


Tests for Q3.2.5 (points: 2)
```
>>> assert p_value1 == sum(statistic1['test statistic'] <= measured_statistic1) / len(statistic1)
```


Tests for Q3.3 (points: 1)
```
>>> assert measured_statistic2 == abs(4 / 60 - 1 / 6)
```


Tests for Q3.4.1. (points: 2)
```
>>> assert len(statistic2) == 1000
>>> assert sum(statistic2['test statistic'] >= 0) == 1000
>>> assert sum(statistic2['test statistic'] <= 1) == 1000
>>> assert sixes.equals(simulated_rolls['roll'] == 6)
>>> assert sim_t2 == abs(sum(sixes) / len(simulated_rolls) - 1 / 6)
```


Tests for Q3.4.3 (points: None)
```
>>> assert p_value2 == sum(statistic2['test statistic'] >= measured_statistic2) / len(statistic2)
```
