```python
# Initialize Otter
import otter
grader = otter.Notebook("lab1.ipynb")
```

# Welcome to Lab 1

Welcome to your first data science lab!
In this lab, we will learn how to use **Jupyter Notebooks** and write our very first lines of Python code.  
No programming experience? No problem! We’ll go step by step. We'll spend the next few weeks explaining Python programming. For now, this is just to whet your appetite and get you familiar with the tools we are using in this class.

**Goals for Today:**
1. Meet your classmates.
2. Learn how to run code in Jupyter.
3. Practice writing and running simple Python commands and Markdown text.

## 1. Hello Team

If you haven't already, take a few minutes to introduce yourselves:

1. Your name  
2. Your area or planned area of study at UNC  
3. Why you are interested in data science  
4. What kind of group project topics you are interested in  

**Bonus question:** If you were a vegetable, which vegetable would you be **and why**?  
(Remember: This is not the same as "What is your favorite vegetable?")


## 2. Hello Jupyter

This is a Jupyter Notebook, which is an interactive “notebook” where we will be writing a mixture of Python code and text explanations of our results. Each of these boxes is a cell, which can contain a piece of code or text. You’ll become very familiar with it over the semester.

- **Markdown cells** (like this one) contain text, which can be prettily formatted.   
- **Code cells** contain Python code you can run.

Let's start by running code that’s already written for you.

Below is a code cell. To run it:  
- Click the cell.
- Press **Shift + Enter** (Windows/Mac) or use the **Run ▶** button in the toolbar.

You should see some text appear *below* the cell.

```python
print("Hello, Data Science!")
```

### Now You Try

In the code cell below, change the text inside the quotation marks (`" "`), and then run it again.

Example:  
```python
print("I love cats and pizza!")
```


```python
print("Replace this with your own message")
```

## 3. Variables in Python

A **variable** stores information we can use later.  
In the example below, we store the name of a vegetable in a variable called `vegetable`.

Run the code cell below to see it in action.


```python
vegetable = "broccoli"
print("You chose:", vegetable)
```

### Now change the variable

Change `"broccoli"` to your chosen vegetable from the icebreaker question, then run the cell again.


```python
vegetable = "broccoli"  # Change this to your vegetable inside the quotation marks like you did before!
print("You chose:", vegetable)
```

## What is a Comment in Code?

A **comment** is text in your code that Python ignores when running the program.  
Comments are for **humans to read**, not for the computer.  
You can use them to explain what your code does, leave notes for yourself, or make the code easier to understand for others.

- In Python, a comment starts with a **`#`** symbol.
- Anything after `#` on the same line will not be executed.

**Example:**
```python
# This line is a comment. It will not run.
print("Hello, world!")  # This prints a greeting


## 4. Adding New Cells in Jupyter

Sometimes you’ll want to add your own cell, either to write more code or add text notes.

**To add a new cell:**
1. Click on an existing cell.
2. Press **`+`** to insert a cell below.  
3. Change the cell type to **Code** (in the toolbar) if it isn’t already.
4. Type your Python code and run it.

**To add a new markdown cell:**
1. Click on an existing cell.
2. Press **`+`** to add a new cell.
3. Change the cell type to **Markdown** (from the dropdown in the toolbar).
4. Type your text. You can use `#` for headings, `**bold**` for bold, and `*italics*` for italics. This cheatsheet (https://commonmark.org/help/) contains a useful and concise set of instructions for Markdown editing.
5. Run the cell to see your formatted text.


### Exercise: Add Your Own Code Cell

1. Add a **new code cell** below this one.
2. In your new cell, write a Python command that prints one of your groupmate's name and their vegatable.  
   Example:  
   ```python
   print("Alex is in my group and he would be a corn.")
   ```
3. Run your cell to see the output.


### Exercise: Add Your Own Markdown Cell

1. Add a **new markdown cell** below this one.
2. In your markdown cell, type:
   - A level-2 heading with your name (`## My Name`)
   - A numbered list of the names of your groupmates.
   - A sentence on why you are excited to learn about data science.
3. Run the cell to see your formatted text.


## 5. Submission to Gradescope

### Step 1 – Run All Cells and Save
- My favorite thing to do in Jupyter Notebook is to click the ⏩️ button on the toolbar. This refreshes the memory, then runs all cells in order.
- Then, click **File → Save Notebook** to make sure your notebook file (`.ipynb`) is up to date. Or run `command (or control) + s`.

### Step 2 – Export to PDF
  1. Choose **File → Print**.
  2. In your browser’s print dialog, change the **Destination** to “Save as PDF.”
  3. Click **Save** and name the file clearly (e.g., `lab1.pdf`).

### Step 3 - Run the Last Cell
1. Run the last cell (below)
2. Download the zip file to your computer.

### Step 4 – Submit to Gradescope
- Log into **Gradescope** and open the Lab 1 assignment. You can get there from Canvas assignment as well.
- Upload **both files** by highlighting both items and draggin and dropping them to the submission box:
  1. Your **`.zip` file** (generated by otter).
  2. Your **`.pdf` file** (exported from the notebook).
- Make sure your submission is **complete and readable** before clicking submit.

### Step 5 – Confirm Your Submission
- After submitting, check Gradescope’s preview to ensure your files display properly.
- If anything is missing or unreadable, re-export and resubmit before the deadline.


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit zip file and PDF to Gradescope Lab1

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for otter-demo (points: 0)
```
>>> assert isinstance(vegetable, str)
```
